```python

```

# Laboratory Exercise 2 - Starter (2 points)

In this assignment, you will build a model to predict whether a customer canceled their reservation using the provided `hotel_booking.csv` dataset.

Your tasks are as follows:

1. Perform Exploratory Data Analysis (EDA) to understand the dataset.
2. Preprocess the data to prepare it for modeling.
3. Train and evaluate a machine learning model to predict  whether a customer canceled their reservation.

**Notes:**

- Some parts of the code are already provided. **Do not modify the existing code.**
- **Do not hard-code** any values or test data. There are hidden tests that will check your implementation.
- Write your solution only in the sections marked with `### YOUR SOLUTION`.
- Each task description indicates whether it is **automatically** or **manually graded**.
- You can verify automatically graded tasks using the cell labeled `### TEST` after each function.

***IMPORTANT NOTE:***
- Name your Jupyter Notebook as `lab_ex_2_{index}.ipynb`.
- For example, if your index is 123456, you should name your notebook as `lab_ex_2_12346.ipynb`.

```python
import os
import hashlib
import random
import numpy as np
import pandas as pd
import seaborn as sns
import matplotlib.pyplot as plt
from sklearn.impute import KNNImputer
from sklearn.preprocessing import OrdinalEncoder, LabelEncoder
from sklearn.model_selection import train_test_split
```

```python
os.environ["PYTHONHASHSEED"] = "42"
random.seed(42)
np.random.seed(42)

RANDOM_STATE = 42
```

```python
def hash_series(series):
    series_str = ",".join(map(str, series.values))
    return hashlib.sha256(series_str.encode()).hexdigest()
```

```python
def hash_data_frame(df):
    df_sorted = df.sort_index(axis=1).sort_values(by=list(df.columns))
    return hashlib.sha256(pd.util.hash_pandas_object(df_sorted, index=True).values).hexdigest()
```

```python
df = pd.read_csv("hotel_bookings.csv")
```

```python
df.head()
```

<!-- grade_id: cell-08b353283c0d0c45, points: 1, grade: true, solution: true -->
```python
### MANUALLY GRADED TASK
def calculate_descriptive_statistics(df):
    """
    Compute the descriptive statistics for all numeric columns in the dataset.
    The statistics should include: count, mean, standard deviation (std), minimum,
    25th percentile, median, 75th percentile, and maximum.
    
    Return the result as a `pd.DataFrame.
    """

    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
calculate_descriptive_statistics(df)
```

<!-- grade_id: cell-cd32f44f46abab58, points: 1, grade: true, solution: true -->
```python
### MANUALLY GRADED TASK
def plot_pairwise_relationships(df):
    """
    Create pairwise plots to visualize the relationships between all numerical variables
    in the dataset using scatter plots and histograms.
    """

    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
plot_pairwise_relationships(df)
```

<!-- grade_id: cell-093bac8eace2bc28, points: 1, grade: true, solution: true -->
```python
### MANUALLY GRADED TASK
def plot_correlation_intensity(df):
    """
    Plot a heatmap showing the strength and direction of the Pearson correlation
    between all numeric variables in the dataset.
    """

    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
plot_correlation_intensity(df)
```

<!-- grade_id: cell-1e69ab9c2a590e52, points: 1, grade: true, solution: true -->
```python
### MANUALLY GRADED TASK
def plot_lead_time_boxplot(df):
    """
    Create a boxplot showing how lead time differs between canceled and non-canceled bookings.
    """

    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
plot_lead_time_boxplot(df)
```

<!-- grade_id: cell-b409bd2f7c5dff66, points: 1, grade: true, solution: true -->
```python
### MANUALLY GRADED TASK
def plot_cancellation_counts(df):
    """
    Create a bar plot showing how many bookings were canceled vs not canceled.
    """

    # YOUR CODE HERE
    raise NotImplementedError()

```

```python
plot_cancellation_counts(df)
```

<!-- grade_id: cell-b1384a94e54f911b, points: 0, grade: false, solution: true -->
```python
### AUTOMATICALLY GRADED TASK
def impute_children_mode(df):
    """
    Fill missing values in `children` with the mode. Use the fillna() method.

    Return the updated dataset as `pd.DataFrame`.
    """
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return df
```

```python
impute_children_mode(df)
```

<!-- grade_id: cell-35295798fe126d3b, points: 1, grade: true, solution: false -->
```python
### TEST
_df = pd.read_csv("hotel_bookings.csv")
_df = impute_children_mode(_df)
assert hash_data_frame(_df) == "2f1c5dfe0828d78ea1a8f9c5976821a9a3eaca07cfe98326304d5972c11b361b"
```

<!-- grade_id: cell-2143289629fae90a, points: 0, grade: false, solution: true -->
```python
### AUTOMATICALLY GRADED TASK
def impute_country_mode(df):
    """
    Fill missing values in `country` with the mode. Use the fillna() method.

    Return the updated dataset as `pd.DataFrame`.
    """
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return df
```

```python
impute_country_mode(df)
```

<!-- grade_id: cell-7460131a7cb16ffd, points: 1, grade: true, solution: false -->
```python
### TEST
_df = pd.read_csv("hotel_bookings.csv")
_df = impute_country_mode(_df)
assert hash_data_frame(_df) == "528991530f1e5408a3bcbcf681a1c62189e516366e6a767ce163320e705487e5"
```

<!-- grade_id: cell-aec99551f95d3552, points: 1, grade: true, solution: true -->
```python
### MANUALLY GRADED TASK
def impute_agent_company_knn(df):
    """
    Impute missing values in `agent` and `company` using KNNImputer with the appropriate columns.
    
    Note: The imputed values should be converted to integer.

    Return the updated dataset as `pd.DataFrame`.
    """
    
    # YOUR CODE HERE
    raise NotImplementedError()

    return df

```

```python
impute_agent_company_knn(df)
```

<!-- grade_id: cell-20d354eb10f4744f, points: 0, grade: false, solution: true -->
```python
### AUTOMATICALLY GRADED TASK
def encode_or_drop_hotel(df):
    """
    Encode the `hotel` variable or remove it from the dataset.

    Note: If you plan to perform one-hot encoding, use the `pd.get_dummies` function and
    append the new columns to the dataset. Also, remove the original column.
    
    Return the updated dataset as `pd.DataFrame`.
    """

    # YOUR CODE HERE
    raise NotImplementedError()
    
    return df
```

```python
encode_or_drop_hotel(df)
```

<!-- grade_id: cell-86478450526f5ee0, points: 1, grade: true, solution: false -->
```python
### TEST
_df = pd.read_csv("hotel_bookings.csv")
_df = encode_or_drop_hotel(_df)
assert hash_data_frame(_df) == "ddfef26391e08ec0e788486481c885ebeecda826fa125e665e4eeea6096b91c9"
```

<!-- grade_id: cell-8e90ad2fdb7ace81, points: 0, grade: false, solution: true -->
```python
### AUTOMATICALLY GRADED TASK
def encode_is_canceled(df):
    """
    Encode the `is_canceled` column manually:
    - "no"  -> 0
    - "yes" -> 1
    
    Return the updated dataset as `pd.DataFrame`.
    """

    # YOUR CODE HERE
    raise NotImplementedError()

    return df
```

```python
encode_is_canceled(df)
```

<!-- grade_id: cell-ae81c002219756cb, points: 1, grade: true, solution: false -->
```python
### TEST
_df = pd.read_csv("hotel_bookings.csv")
_df = encode_is_canceled(_df)
assert hash_data_frame(_df) == "bd5c6cf86d2ef9dc64162cb8d9d8be54b1247e6a0da08074085d6c4bf86b71df"
```

<!-- grade_id: cell-709ca8c88561ea91, points: 0, grade: false, solution: true -->
```python
### AUTOMATICALLY GRADED TASK
def encode_arrival_date_month(df):
    """
    Encode the `arrival_date_month` column using OrdinalEncoder
    with fixed calendar order (January → December).

    Note: The encoded values should be converted to integer.

    Return the updated dataset as `pd.DataFrame`.
    """
    
    # YOUR CODE HERE
    raise NotImplementedError()

    return df

```

```python
encode_arrival_date_month(df)
```

<!-- grade_id: cell-5a360390b24615c8, points: 1, grade: true, solution: false -->
```python
### TEST
_df = pd.read_csv("hotel_bookings.csv")
_df = encode_arrival_date_month(_df)
assert hash_data_frame(_df) == "5def409216f31ceb05b96377d444141c7e28cfca808a8481484ab0443a873c99"
```

<!-- grade_id: cell-292db53c326a47c2, points: 0, grade: false, solution: true -->
```python
### AUTOMATICALLY GRADED TASK
def encode_or_drop_meal(df):
    """
    Encode the `meal` variable or remove it from the dataset.

    Note: If you plan to encode it use OrdinalEncoder with automatically detected categories. 
    The encoded values should be converted to integer.
    
    Return the updated dataset as `pd.DataFrame`.
    """

    # YOUR CODE HERE
    raise NotImplementedError()

    return df
```

```python
encode_or_drop_meal(df)
```

<!-- grade_id: cell-c06555f2b151deb3, points: 1, grade: true, solution: false -->
```python
### TEST
_df = pd.read_csv("hotel_bookings.csv")
_df = encode_or_drop_meal(_df)
assert hash_data_frame(_df) == "1c7fbbbcd1b083a264a946bce4188949dd3d27882ce357b27e5ecddee7c3c3e2"
```

<!-- grade_id: cell-c82ce5337a541449, points: 0, grade: false, solution: true -->
```python
### AUTOMATICALLY GRADED TASK
def encode_or_drop_country(df):
    """
    Encode the `country` variable or remove it from the dataset.

    Note: If you plan to encode it use OrdinalEncoder with automatically detected categories. 
    The encoded values should be converted to integer.
    
    Return the updated dataset as `pd.DataFrame`.
    """

    # YOUR CODE HERE
    raise NotImplementedError()

    return df
```

```python
encode_or_drop_country(df)
```

<!-- grade_id: cell-a7778219d850a04a, points: 1, grade: true, solution: false -->
```python
### TEST
_df = pd.read_csv("hotel_bookings.csv")
_df = encode_or_drop_country(_df)
print(hash_data_frame(_df))
assert hash_data_frame(_df) == "be2489fa971ddb0bd4043d782985f94e0f65949922867165b475fb553138ce56"
```

<!-- grade_id: cell-23564994e9e8e16d, points: 0, grade: false, solution: true -->
```python
### AUTOMATICALLY GRADED TASK
def encode_or_drop_customer_type(df):
    """
    Encode the `customer_type` variable or remove it from the dataset.

    Note: If you plan to encode it use OrdinalEncoder with automatically detected categories. 
    The encoded values should be converted to integer.
    
    Return the updated dataset as `pd.DataFrame`.
    """

    # YOUR CODE HERE
    raise NotImplementedError()

    return df
```

```python
encode_or_drop_customer_type(df)
```

<!-- grade_id: cell-856d15911bbe3a05, points: 1, grade: true, solution: false -->
```python
### TEST
_df = pd.read_csv("hotel_bookings.csv")
_df = encode_or_drop_customer_type(_df)
assert hash_data_frame(_df) == "6c55892a0f98882ac8720b5ec0350a71cfb936ebcde9472c604ce364ce1da6a5"
```

<!-- grade_id: cell-dcd4d883f258559d, points: 0, grade: false, solution: true -->
```python
### AUTOMATICALLY GRADED TASK
def encode_deposit_type(df):
    """
    Encode the `deposit_type` column using OrdinalEncoder
    with fixed order ['No Deposit', 'Non Refund', 'Refundable'].

    Note: The encoded values should be converted to integer.

    Return the updated dataset as `pd.DataFrame`.
    """
    
    # YOUR CODE HERE
    raise NotImplementedError()

    return df
```

```python
encode_deposit_type(df)
```

<!-- grade_id: cell-f540362c84cada67, points: 1, grade: true, solution: false -->
```python
### TEST
_df = pd.read_csv("hotel_bookings.csv")
_df = encode_deposit_type(_df)
assert hash_data_frame(_df) == "d580508ad442d05888f34e34be95f40be99758e97d57cf98bc5085ff48df32ba"
```

<!-- grade_id: cell-fad0840db50f8814, points: 0, grade: false, solution: true -->
```python
### AUTOMATICALLY GRADED TASK
def combine_person_columns(df):
    """
    Create a new column 'total_persons' as the sum of 'adults', 'children', and 'babies'.
    Then drop the original three columns.

    Return the updated dataset as `pd.DataFrame`.
    """
    
    # YOUR CODE HERE
    raise NotImplementedError()

    return df
```

```python
combine_person_columns(df)
```

<!-- grade_id: cell-20d6569238ffd765, points: 1, grade: true, solution: false -->
```python
### TEST
_df = pd.read_csv("hotel_bookings.csv")
_df = combine_person_columns(_df)
assert hash_data_frame(_df) == "f0d93e268feae1e4416b8f414b214844f2f265115d1a6a9c867586278d94b859"
```

<!-- grade_id: cell-12451dc4821c039b, points: 0, grade: false, solution: true -->
```python
### AUTOMATICALLY GRADED TASK
def create_room_match_column(df):
    """
    Create a new Boolean column 'room_match' that indicates whether
    the reserved room type matches the assigned room type.
    Then drop the original two columns.

    Return the updated DataFrame.
    """
    # YOUR CODE HERE
    raise NotImplementedError()

    return df

```

```python
create_room_match_column(df)
```

<!-- grade_id: cell-a769fe24a6949816, points: 1, grade: true, solution: false -->
```python
### TEST
_df = pd.read_csv("hotel_bookings.csv")
_df = create_room_match_column(_df)
assert hash_data_frame(_df) == "4d99349e3d7901bb36e2752e8eb34a42d9a84fc4dba6dec68f6435111985efef"
```

<!-- grade_id: cell-6d1422fd9c41c49f, points: 0, grade: false, solution: true -->
```python
### AUTOMATICALLY GRADED TASK
def partition_dataset(df):
    """
    Split the dataset into features (X) and target (y), where the target is `Price (Euro)`.
    Then, divide X and y into training and test sets using an 80:20 ratio.

    Use `random_state=RANDOM_STATE` to ensure reproducibility.
    
    Return the sets in the following order: train_X, test_X, train_y, test_y.
    """

    # YOUR CODE HERE
    raise NotImplementedError()
    
    return train_X, test_X, train_y, test_y
```

```python
train_X, test_X, train_y, test_y = partition_dataset(df)
```

<!-- grade_id: cell-8de6535f967cc8ef, points: 1, grade: true, solution: false -->
```python
_df = pd.read_csv("hotel_bookings.csv")
_train_X, _test_X, _train_y, _test_y = partition_dataset(_df)
assert hash_data_frame(_train_X) == "9ac702152f9fbb53b83066adf10226413d2e61a67941fa6297335ef1912e0523"
assert hash_data_frame(_test_X) == "2f646b11dd0038298d3c0e990a795266f1deb2efa996ab2eb5fa40a833167b79"
assert hash_series(_train_y) == "97d345bcc07a58c45d8353e65862fb370191c4b7e6c9dd9c99eb3423f943977f"
assert hash_series(_test_y) == "a043cd1df6fd8349f79098436bdc8ccd4df94ae494784fab62a4716044dbea53"
```

<!-- grade_id: cell-777bdf9235f5d8a2, points: 1, grade: true, solution: true -->
```python
### MANUALLY GRADED TASK
def fit_model(train_X, train_y):
    """
    Train a suitable model.
    
    Return the trained model.
    """

    # YOUR CODE HERE
    raise NotImplementedError()
    
    return model
```

```python
model = fit_model(train_X, train_y)
```

```python
pred_y = model.predict(test_X)
```

<!-- grade_id: cell-df42bb5e36801efb, points: 1, grade: true, solution: true -->
```python
### MANUALLY GRADED TASK
def evaluate_model(test_y, pred_y):
    """
    Evaluate the model's predictions using different metrics.
    
    Print the results.
    """

    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
evaluate_model(test_y, pred_y)
```