# Laboratory Exercise 1 - Starter (2 points)

In this assignment, you will perform **web scraping** on the website [https://books.toscrape.com/](https://books.toscrape.com/).

Your tasks are as follows:

1. **Extract** `image_url`, `rating`, `title`, and `price` for the 20 books displayed on the main page.
2. **Organize the data** into a **pandas.DataFrame**.
3. **Visualize** the collected data according to the given requirements.

**Notes:**

- Some parts of the code are already provided. **Do not modify the existing code.**
- **Do not hard-code** any values or test data. There are hidden tests that will check your implementation.
- Write your solution only in the sections marked with `### YOUR SOLUTION`.
- Each task description indicates whether it is **automatically** or **manually graded**.
- You can verify automatically graded tasks using the cell labeled `### TEST` after each function.

***IMPORTANT NOTE:***
- Name your Jupyter Notebook as `lab_ex_1_{index}.ipynb`.
- For example, if your index is 123456, you should name your notebook as `lab_ex_1_12346.ipynb`.

```python
import requests
import pandas as pd
import matplotlib.pyplot as plt
import seaborn as sns
from bs4 import BeautifulSoup
```

```python
url = "https://books.toscrape.com/"
```

```python
response = requests.get(url)
```

```python
page_element = BeautifulSoup(response.text, "html.parser")
```

```python
product_elements = page_element.select("article.product_pod")
```

**(0.2 points)** Implement the `get_image_url` function to return the image URL of a given product element using **BeautifulSoup**. For example, the first product on the page has the image URL `"media/cache/2c/da/2cdad67c44b002e7ead0cc35693c0e8b.jpg"`.

<!-- grade_id: cell-2916ab1737f17d0b, points: 0, grade: false, solution: true -->
```python
### AUTOMATICALLY GRADED TASK
def get_image_url(product_element):
    "Get the URL of the image from the given product element using BeautifulSoup."

    # YOUR CODE HERE
    raise NotImplementedError()
```

Test your implementation by running the following cell.

<!-- grade_id: cell-30a6648b0f0ec9ee, points: 0.2, grade: true, solution: false -->
```python
### TEST
assert get_image_url(product_elements[0]) == "media/cache/2c/da/2cdad67c44b002e7ead0cc35693c0e8b.jpg"
```

**(0.2 points)** Implement the `get_rating` function to return the product’s rating as an integer between 1 and 5. For example, the first product on the page has a rating of `3`.

<!-- grade_id: cell-1263d0f20bfd0ed0, points: 0, grade: false, solution: true -->
```python
### AUTOMATICALLY GRADED TASK
def get_rating(product_element):
    "Get the rating (1-5) of the given product element."

    # YOUR CODE HERE
    raise NotImplementedError()
```

Test your implementation by running the following cell.

<!-- grade_id: cell-9b6c7f9737451042, points: 0.2, grade: true, solution: false -->
```python
### TEST
assert get_rating(product_elements[0]) == 3
```

**(0.2 points)** Implement the `get_title` function to return the full title of the given product as a string. For example, the first product on the page has the title `"A Light in the Attic"`.

<!-- grade_id: cell-bfed43972258ef45, points: 0, grade: false, solution: true -->
```python
### AUTOMATICALLY GRADED TASK
def get_title(product_element):
    "Get the full-length title of the given product element."

    # YOUR CODE HERE
    raise NotImplementedError()
```

Test your implementation by running the following cell.

<!-- grade_id: cell-cee445e198fed068, points: 0.2, grade: true, solution: false -->
```python
### TEST
assert get_title(product_elements[0]) == "A Light in the Attic"
```

**(0.2 points)** Implement the `get_price` function to return the product’s price as a float. For example, the first product on the page has a price of `51.77`.

<!-- grade_id: cell-d4da8aaa77639d69, points: 0, grade: false, solution: true -->
```python
def get_price(product_element):
    "Get the price of the given product element."

    # YOUR CODE HERE
    raise NotImplementedError()
```

Test your implementation by running the following cell.

<!-- grade_id: cell-94f30d465d40cf21, points: 0.2, grade: true, solution: false -->
```python
### TEST
assert get_price(product_elements[0]) == 51.77
```

```python
def get_attributes(product_element):
    return {
        "image_url": get_image_url(product_element),
        "rating": get_rating(product_element),
        "title": get_title(product_element),
        "price": get_price(product_element),
    }
```

```python
get_attributes(product_elements[0])
```

```python
product_data = [get_attributes(product_element) for product_element in product_elements]
```

```python
df = pd.DataFrame(product_data)
```

```python
df.head()
```

***If you were not able to complete the previous tasks successfully, you may uncomment the following line to continue with the next steps.***

```python
# df = pd.read_csv("products.csv")
```

**(0.2 points)** Create a visualization showing the number of products for each rating (1-5).

<!-- grade_id: cell-fc2d7650b3805138, points: 0.2, grade: true, solution: true -->
```python
### MANUALLY GRADED TASK
def plot_counts_by_rating(df):
    "Plot the count of products by rating."""

    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
plot_counts_by_rating(df)
```

**(0.2 points)** Create a visualization showing the distribution of product prices.

<!-- grade_id: cell-50327b69fe70d25e, points: 0.2, grade: true, solution: true -->
```python
### MANUALLY GRADED TASK
def plot_price_distribution(df):
    """Plot the price distribution of the products."""

    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
plot_price_distribution(df)
```

**(0.2 points)** Create a visualization showing the distribution of product title lengths.

<!-- grade_id: cell-e9dc38d2f366378e, points: 0.2, grade: true, solution: true -->
```python
### MANUALLY GRADED TASK
def plot_title_length_distribution(df):
    """Plot the distribution of the title lengths."""

    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
plot_title_length_distribution(df)
```

**(0.2 points)** Create box plots showing the distribution of product prices for each rating.

<!-- grade_id: cell-166144cfad577966, points: 0.2, grade: true, solution: true -->
```python
### MANUALLY GRADED ANSWER
def plot_price_distribution_per_rating(df):
    """Plot the price distribution of the products per rating."""

    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
plot_price_distribution_per_rating(df)
```

**(0.2 points)** Create a heatmap showing the Pearson correlation between the numeric variables in the dataset.

<!-- grade_id: cell-9be6a485a8e959e3, points: 0.2, grade: true, solution: true -->
```python
### MANUALLY GRADED ANSWER
def plot_corr_heat_map(df):
    """Plot the Pearson correlation between the numeric variables in the dataset using a heat map."""

    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
plot_corr_heat_map(df)
```

**(0.2 points)** Compute the `count`, `mean`, `std`, `min`, `25%`, `median`, `75%`, and `max` statistics of product prices grouped by rating (1-5).

<!-- grade_id: cell-9dea05ecfb980c40, points: 0.2, grade: true, solution: true -->
```python
### MANUALLY GRADED ANSWER
def compute_descriptive_stats_for_price_per_rating(df):
    """
    Compute the `count`, `mean`, `std`, `min`, `25% quantile`, `median`, `75% quantile`, and `max` statistics
    for the product's prices per rating (1-5). Return them as a `pandas.DataFrame`.
    """

    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
compute_descriptive_stats_for_price_per_rating(df)
```